```python
# Initialize Otter
import otter
grader = otter.Notebook("class_wed.ipynb")
```

> **About the autograder:** This notebook uses [otter-grader](https://otter-grader.readthedocs.io/) for instant feedback. After completing each task, run the `grader.check("t1")` cell below it — a ✅ means your answer is correct. This exercise is **not graded**; the checks are just for your own feedback. If a check fails, read the message and try again.

# Week 1 — Wednesday: Python Review for Machine Learning

**DATA 202 · Calvin University**

Work through each section with a partner: read the explanation, run the examples, then complete the task together.

> *What do we need to remember about Python to get started with machine learning?*

1. Lists and list comprehensions
2. Tuples
3. Dictionaries
4. Objects: attributes and methods
5. String manipulation
6. NumPy arrays

---
## 1. Lists and List Comprehensions

Lists are **ordered, mutable** collections — you can change them after creation.

```python
nums = [1, 2, 3, 4]
print(nums[0])      # first element
nums.append(5)      # add to end
nums[2] = 10        # replace an element
print(nums)
```

**List comprehensions** are a concise way to build a new list from an existing one:

```python
squares = [x**2 for x in range(6)]
print(squares)

evens = [x for x in range(10) if x % 2 == 0]
print(evens)
```

### Task 1

1. Create a list `temps_c = [0, 20, 37, 100]` (temperatures in Celsius).
2. Use a list comprehension to convert it to Fahrenheit: `f = c * 9/5 + 32`. Assign the result to `temps_f`.
3. Print the Fahrenheit list.

> **Discuss:** did any step modify `temps_c` in place, or create a new list?

```python
temps_c = ...
temps_f = ...
print(temps_f)
```

```python
grader.check("t1")
```

---
## 2. Tuples

Tuples are **ordered but immutable** — you cannot change them after creation. Commonly used to return multiple values from a function.

```python
point = (3, 4)
x, y = point          # tuple unpacking
print(f'x={x}, y={y}')

def min_max(values):
    return (min(values), max(values))

lo, hi = min_max([5, 2, 9, 1])
print(f'min={lo}, max={hi}')
```

### Task 2

1. Write a function `summary(values)` that returns a tuple `(min, max, mean)`. Hint: `sum(values) / len(values)` gives the mean.
2. Call it with `[10, 20, 30, 40, 50]` and unpack into `lo, hi, avg`.
3. Print: `Range: 10-50, average: 30.0`

> **Discuss:** why is a tuple more appropriate here than a list?

```python
...
    ...

lo, hi, avg = ...
print(f'Range: {lo}-{hi}, average: {avg}')
```

```python
grader.check("t2")
```

---
## 3. Dictionaries

Dictionaries store **key-value pairs** with fast lookups. In ML you will encounter them constantly: model parameters, configuration, label encodings, evaluation metrics.

```python
student = {'name': 'Alice', 'age': 20}
print(student['age'])       # access by key
student['grade'] = 'A'      # add a new key

for key, value in student.items():
    print(f'{key}: {value}')
```

### Task 3

1. Create `label_map = {'low': 0, 'medium': 1, 'high': 2}`.
2. Given `raw = ['high', 'low', 'medium', 'high', 'low']`, use a list comprehension to produce numeric codes. Assign to `codes`.
3. Print the result: `[2, 0, 1, 2, 0]`

> **Discuss:** where else in a data pipeline might you use a mapping like this?

```python
label_map = ...
raw = ...
codes = ...
print(codes)
```

```python
grader.check("t3")
```

---
## 4. Objects: Attributes and Methods

Almost everything in Python is an **object** with attributes (stored values) and methods (functions attached to it).

- `object.method()` — called *on* the object
- `function(object)` — called *with* the object as argument

Watch for **mutability**: some methods modify in place, others return a new object.

```python
s = 'hello'
print(s.upper())    # method — returns NEW string (strings are immutable)
print(s.islower())  # method — returns True/False
print(len(s))       # function, not method
print(s)            # s is still 'hello' — nothing changed it
```

```python
nums = [3, 1, 4, 1, 5]
nums.sort()                   # modifies IN PLACE
print('after sort():', nums)

nums2 = [3, 1, 4]
sorted_nums = sorted(nums2)   # returns a NEW list
print('nums2 unchanged:', nums2)
print('sorted_nums:', sorted_nums)
```

### Task 4

1. Given `words = ['Machine', 'learning', 'IS', 'fascinating']`, use a list comprehension with a string method to lowercase every word. Assign to `lower_words`.
2. Sort the result with `sorted()` and store it in `sorted_words`.
3. Print both `words` and `sorted_words` — confirm `words` is unchanged.

> **Discuss:** why does this distinction (in-place vs. new object) matter in a data pipeline?

```python
words = ['Machine', 'learning', 'IS', 'fascinating']
lower_words = ...
sorted_words = ...
print(words)
print(sorted_words)
```

```python
grader.check("t4")
```

---
## 5. String Manipulation

Real datasets are messy. Text columns almost always need cleaning before a model can use them.

```python
text = '  data,science,rocks!  '
print(text.strip())                    # remove surrounding whitespace
print(text.strip().split(','))         # split into a list
print('-'.join(['data', 'science']))   # join with a separator
print(text.replace('rocks', 'rules'))  # find and replace
```

### Task 5

You receive a raw string from a form submission:

```python
entry = "  Grand Rapids ; Michigan ; 49546  "
```

1. Strip leading/trailing whitespace, then split by `";"` to get a list of parts.
2. Strip whitespace from each part using a list comprehension.
3. Build a dictionary `location = {'city': ..., 'state': ..., 'zip': ...}` from the parts.

> **Discuss:** why would inconsistent formatting (extra spaces, different separators) break downstream analysis?

```python
entry = "  Grand Rapids ; Michigan ; 49546  "
parts = ...
location = ...
print(location)
```

```python
grader.check("t5")
```

---
## 6. NumPy Arrays

Python lists are flexible but not designed for numerical computing:

```python
[1, 2, 3] + [4, 5, 6]  # concatenation, NOT addition
[1, 2, 3] * 2           # repetition, NOT multiplication
```

**NumPy arrays** apply operations **element-wise** and are much faster.

```python
import numpy as np

arr = np.array([1, 2, 3, 4, 5])
print(arr * 2)                          # [2 4 6 8 10]
print(arr + np.array([10,20,30,40,50])) # element-wise add
```

```python
print(arr.sum(), arr.mean(), arr.std())
```

```python
# 2D arrays: rows = samples, columns = features
mat = np.array([[1, 2, 3],
                [4, 5, 6]])
print(mat.shape)   # (2, 3)
print(mat[0, 1])   # row 0, col 1 -> 2
print(mat[:, 1])   # all rows, column 1 -> [2 5]
```

### Task 6

You have exam scores for 4 students across 3 assignments:

```python
scores = np.array([[88, 92, 95],
                   [78, 85, 80],
                   [90, 91, 89],
                   [70, 72, 68]])
```

1. Compute the average score per student (mean across each row, `axis=1`). Assign to `avg_per_student`.
2. Compute the average score per assignment (mean down each column, `axis=0`). Assign to `avg_per_assignment`.
3. Find the index of the student with the highest overall average using `np.argmax()`. Assign to `best_student`.

> **Discuss:** how does this 2D structure relate to how ML models see training data?

```python
import numpy as np

scores = np.array([[88, 92, 95],
                   [78, 85, 80],
                   [90, 91, 89],
                   [70, 72, 68]])

avg_per_student = ...
avg_per_assignment = ...
best_student = ...
print('Avg per student:', avg_per_student)
print('Avg per assignment:', avg_per_assignment)
print('Best student index:', best_student)
```

```python
grader.check("t6")
```

---
## Wrap-Up

These building blocks appear constantly in the ML stack:

| Concept | Where you will see it |
|---|---|
| Lists / comprehensions | Feature lists, label encoding, data pipelines |
| Tuples | Train/test splits, function returns, coordinate pairs |
| Dictionaries | Model parameters, configuration, label mappings |
| Methods vs. functions | `.fit()`, `.transform()`, `.predict()` in scikit-learn |
| String manipulation | Cleaning text columns before analysis |
| NumPy arrays | Feature matrices, model inputs/outputs, vectorized math |

Tests for t1 (points: 1)
```
>>> assert temps_c == [0, 20, 37, 100], 'temps_c should be unchanged'
>>> assert temps_f == [32.0, 68.0, 98.6, 212.0], 'check your Fahrenheit formula'
```


Tests for t2 (points: 1)
```
>>> assert lo == 10
>>> assert hi == 50
>>> assert avg == 30.0
```


Tests for t3 (points: 1)
```
>>> assert codes == [2, 0, 1, 2, 0]
```


Tests for t4 (points: 1)
```
>>> assert words == ['Machine', 'learning', 'IS', 'fascinating'], 'words should be unchanged'
>>> assert lower_words == ['machine', 'learning', 'is', 'fascinating']
>>> assert sorted_words == ['fascinating', 'is', 'learning', 'machine']
```


Tests for t5 (points: 1)
```
>>> assert location['city'] == 'Grand Rapids'
>>> assert location['state'] == 'Michigan'
>>> assert location['zip'] == '49546'
```


Tests for t6 (points: 1)
```
>>> assert avg_per_student.shape == (4,)
>>> assert avg_per_assignment.shape == (3,)
>>> assert best_student == 2, 'student index 2 has the highest average'
```
